# EC Notebook for Lecture 11: Bar Graphs vs. Histograms

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Monday, Sept. 23** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "hello" dataset -- as we did in lecture 11 -- to plot a few graphs.

Using your @illinois.edu address, you can access the "hello" dataset here: 
https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0

Always, remember to load your data first.

<!-- grade_id: cell-1628b09530950d23, points: 0, grade: false, solution: true -->
```python

```

### Question 1:

Today we will consider non-freshman students. First, find the mean of number of hours a non-freshman student sleeps (hmean_non_freshman). Next, plot the histogram for number of hours of sleep of non-freshman students. 

(Note that you should expect hmean_non_freshman value to have highest bar (Right?)!)

<!-- grade_id: cell-99d959c1fa5c1922, points: 0, grade: false, solution: true -->
```python
hmean_non_freshman = ...
print(hmean_non_freshman)
#Now, write the code for the plot



```

### Question 2: 

Next, add 2 to number of hours of sleep of Junior students only. Then, plot a histogram for number of hours of sleep of non-freshman students.


(Note that in the class, we added a number to every row and saw that the histogram just shifted towards right. But here we are adding a number to a few rows only.)

### Question 2:

Suppose we are interested in knowing students from which year watch twitch the most. Create a bar plot for number of students who have youtube. 

<!-- grade_id: cell-0f2b9437c5b3e29a, points: 0, grade: false, solution: true -->
```python


```

### Question 3 

Note that in the above plot, the bar may be higher just because there are more students in the class from that year. Calculate the fraction of students from each year who have youtube. Try to use the group you just created and functions related to groups which you learned in the class.

<!-- grade_id: cell-ea88dcf40cc90ada, points: 0, grade: false, solution: true -->
```python
freshman_frac = ...
sophomore_frac = ...
junior_frac = ...
senior_frac = ...
other_frac = ...
print(freshman_frac)
print(sophomore_frac)
print(junior_frac)
print(senior_frac)
print(other_frac)



```

(You don't have to answer this. Just think for yourself!):
What do you think now, shouldn't we have plotted the fraction to gain a better understanding?

<!-- grade_id: cell-357bac434e050bca, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(int(hmean_non_freshman) == 7), "The mean does not appear correct."
assert(round(freshman_frac,2) == 0.84), "The freshman_frac does not appear correct ."
assert(round(sophomore_frac,2) == 0.94), "The sophomore_frac does not appear correct ."
assert(round(junior_frac,2) == 0.84), "The junior_frac does not appear correct ."
assert(round(senior_frac,2) == 0.89), "The senior_frac does not appear correct ."
assert(round(other_frac,2) == 0.86), "The other_frac does not appear correct ."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)

```python

```