# EC Notebook for Lecture 23: The Normal Approximation

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Oct. 23** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## Try `scipy`

In the lecture, we used another library of python to find the areas under the normal curve, which is called `scipy`. You can try to import the library as we did in the lecture, if it is not installed, use the command:

`conda install -c anaconda scipy` 

in your terminal to install it.

### Calculate `CDF` and `PDF`

Using `scipy`, calculate the `cdf` and `pdf`  when `z = 0.8`.


<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
v_cdf = ...
v_pdf = ...


print(v_cdf, v_pdf)
```

## Approximate Normal Distribution

Normal distribution is everywhere in our life. Many histograms are close to the normal curve, now let's take an example involving the rolling of dies.

### Average of Dice
Create a function to roll 50 dices and return the average of numbers you get from these dices.


<!-- grade_id: cell-060c81a4477630d2, points: 0, grade: false, solution: true -->
```python
def avg_dice():
    ...
    

avg_dice()
```

### Distribution of Average
Write a simulation to run the above function for 10,000 times and store the results you get for each time in a DataFrame called `df`.

<!-- grade_id: cell-8b0f0ef67575550d, points: 0, grade: false, solution: true -->
```python
df = ...



# Print out a few random rows:
print(df.sample(5))
# Let’s build some visualizations:
df.hist(bins=20)
```

Does it look like a normal distribution? Now we can calculate the Z-scores to convert data to standard units by the formula you have seen in the lecture.

###  Z-scores
Firstly calculate the average and SD of the average of dices. Then compute the Z-scores of each row and store the values in an extra column called `Z-score`.

<!-- grade_id: cell-9b6af2e58eb5bdd7, points: 0, grade: false, solution: true -->
```python
avg_avg = ...
avg_sd = ...
df["Z-score"] = ...



df["Z-score"].hist(bins=20)
```

Now create another function to calculate the following results with input `value`:
1. `prob`: The probability that the average is less and equal than `value`. 
2. `z_value`: The z-score of value `value`
3. `cdf_value`: value of cdf when `z = z_value`.


<!-- grade_id: cell-24083d8445b4840b, points: 0, grade: false, solution: true -->
```python
def calc_bound(value):
    prob = ...
    z_value = ...
    cdf_value = ...
    print(prob, z_value, cdf_value)


#Check your function! (Hint: the probability should be close to the cdf_value)
for v in [3.3, 3.4, 3.5, 3.6, 3.7]:
    prob, cdf_value = calc_bound(v)
    print (prob, cdf_value)
```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1,2,3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(v_cdf-0.788144)<0.001 and abs(v_pdf-0.2896915)<0.001), "Make sure you calculated the cdf and pdf correctly."
assert(avg_dice() <= 6 and avg_dice() >= 1), "Make sure you calculated the average."
assert(len(df) == 10000), "Make sure you simulated exactly 10,000 times"
for v in [3.3, 3.4, 3.5, 3.6, 3.7]:
    a,b = calc_bound(v)
    assert(abs(a-b)<0.1), "Make sure you calculated the probability correctly."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

```python

```

```python

```