# EC Notebook for Lecture 16: Simulation in Python

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Monday, Oct. 7** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## Accumulating Simulation Data 

In lecture, you saw a six-step rule for accumulating simulation data.  Let's practice one:


### Part 1: Accumulating the data

Tommy, Jack, Michael and Ben are all Illini Football players.  The number of yards they rush each game can be simulated by a random number from 20 to 50 yards.

Write a simulation of football games and store 10,000 observations in a DataFrame called `df`

<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
# Simulation of 10,000 football games:






# Save the simulation data into a DataFrame:
df = ...


# Print out a few random rows:
df.sample(5)
```

### Part 2: Estimate the probability of rushing 140 yards

Estimate the probability that, in a given game, Tommy, Jack, Michael, and Ben rush a total of **at least 140 yards**.  Store your probability estimate in `P_rush140`:

<!-- grade_id: cell-9b6af2e58eb5bdd7, points: 0, grade: false, solution: true -->
```python



P_rush140 = ...
P_rush140


```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1 and Part 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df) == 10000), "Make sure your df has exactly 10,000 observations"
assert(P_rush140 > 0 and P_rush140 < 1), "Make your P_rush140 is a probability"
assert(abs(P_rush140 - 0.5) < 0.1), "Make your P_rush140 estimates the probability of rushing at least 140 yards"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print(f"You found: P(140 yards or more) = {round(100 * P_rush100, 2)}%")
```

### Part 3: Estimate the probability of one or more player rushing more than 45 yards

Estimate the probability that, in a given game, Tommy, Jack, Michael, **or** Ben rush a total of **more than 45 yards in a single game**.  Store your probability estimate in `P_45plus`:

<!-- grade_id: cell-f515099e60a97e4f, points: 0, grade: false, solution: true -->
```python



P_45plus = ...
P_45plus


```

<!-- grade_id: cell-0294c9244f8fe105, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(P_45plus > 0 and P_45plus < 1), "Make your P_45plus is a probability"
assert(abs(P_45plus - 0.4948) < 0.1), "Make your P_45plus estimates the probability of rushing at least one player rushing 45 yards"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print(f"You found: P(at least one player rushing 45 yards) = {round(100 * P_45plus, 2)}%")
print(f"        CHALLENGE: Try to compute this value by hand, it's a bit ticky but a lot of fun! :)")
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)