# EC Notebook for Lecture 14: 

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Monday, Sept. 30** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Probability Words" dataset -- to answer a few questions on probabilities.

You can find this dataset here:
https://docs.google.com/spreadsheets/d/18zQc-cYb376QA478lgWnMbQj39M9wVqHtjbvLP0PTSc/edit#gid=0

<!-- grade_id: cell-f1a1e86012dbeeaf, points: 0, grade: false, solution: true -->
```python
df = ...


```

## Question 1

Let's explore a connection between percentage and probability.

Each row in our "Probability Words" dataset is one person's response to their perception of probability words.  If we selected one from the class at random, what is the probability that their perception of "Almost Certain" is 90% or higher?

*Hint: You may need to use* `len(df)` *to find out how many people rows are in a DataFrame after selecting for a certain condition.*

<!-- grade_id: cell-b21768c2c44e7d30, points: 0, grade: false, solution: true -->
```python
p_almost_certain_90 = ...
p_almost_certain_90


```

<!-- grade_id: cell-e3dc4d39ea8750f4, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(p_almost_certain_90,3)== 0.708)

#
# You will only see this message (with the emoji showing) if you passed all test cases:
#
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Question 2

Now, let's work with random numbers. Remember to ***import random***.

- Let us represent a coin flip as a `0` or `1`.
  * `0` for heads
  * `1` for tails

Randomly generate 10 coin flips in `toss1` through `toss10`:

<!-- grade_id: cell-ee061eb490645b6f, points: 0, grade: false, solution: true -->
```python
# import random first

# Now, coin tossing
toss1 = ...
toss2 = ...
toss3 = ...
toss4 = ...
toss5 = ...
toss6 = ...
toss7 = ...
toss8 = ...
toss9 = ...
toss10 = ...



# Some basic calculations based on your tosses:
heads_ct = toss1 + toss2 + toss3 + toss4 + toss5 + toss6+ toss7 + toss8 + toss9 + toss10
tails_ct = 10 - heads_ct
probability_head = heads_ct / 10

print(f"Number of heads: {heads_ct}")
print(f"Number of tails: {tails_ct}")
print()
print(f"P(head) from your 10 tosses: {probability_head}")
```

## Run it Again

Run the previous cell several times -- watch how the numbers differ! :)

<!-- grade_id: cell-a0e0e65e61139ea8, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert((0 <= probability_head) and (probability_head <= 1)), "The probability is always between 0 and 1, isn't it?"

#
# You will only see this message (with the emoji showing) if you passed all test cases:
#
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Interested in Magic! Run your simulation again and see if the probability of heads changes or not!
What could be the possible reason. You don't have to answer it here, just think about it!

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)