# EC Notebook for Lecture 8: Mean and Standard Deviation

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Monday, Sept. 16** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Hello" dataset -- as we did in lecture 7 -- to answer a few questions on working with columns.

Using your @illinois.edu address, you can access the "Hello" dataset here: https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0

<!-- grade_id: cell-cd5a39351d2266b0, points: 0, grade: false, solution: true -->
```python

```

## 2. Spread of sleeping hours among Freshmen
After coming to college, some people start burning the midnight oil from the very first day and hence sleep less. At the same time, there are few people who are still tired because of the celeberation they had after graduating from high school and hence sleep more. Let's find out what is the average and the spread of sleeping hours among freshmen in our class.

### Question 1: 
Using the "hello.csv" dataset, calculate the mean and the standard deviation (std) in number of hours of sleep among freshmen. 
(Please note that "std" is a common abbreviation for standard deviation.)

<!-- grade_id: cell-9b98c6ac7440756f, points: 0, grade: false, solution: true -->
```python
freshmen_sleep_mean = ...
freshmen_sleep_std = ...
print(freshmen_sleep_mean)
print(freshmen_sleep_std)


```

Now, we are three weeks into the FALL 2019. With all the assignments piling up, our Freshmen friends are having a hard time catching up. Let's reduce their number of hours of sleep (by just 1 hour :) ) so that they don't look tired in Lincoln Hall Theatre (otherwise Wade and Karle will be worried).

### Question 2:
Subtract 1 from "Hours of Sleep" for Freshmen and recalculate the new mean and new standard deviation (std). Notice which quantity changes and which does not on recalculation.

<!-- grade_id: cell-a2a79acc8cdee472, points: 0, grade: false, solution: true -->
```python
freshmen_sleep_new_mean = ...
freshmen_sleep_new_std = ...
print(freshmen_sleep_new_mean)
print(freshmen_sleep_new_std)


```

<!-- grade_id: cell-5ab6337558dbb895, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(int(freshmen_sleep_mean) == 7), "This does not appear correct."
assert(round(freshmen_sleep_std,2) == 0.81), "This does not appear correct ."
assert(int(freshmen_sleep_mean_new) == 6), "This does not appear correct."
assert(round(freshmen_sleep_std_new,2) == 0.81),  "This does not appear correct."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)