# Your grades and GPA at Illinois!

In this lab, you will explore the GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

## Part 0 | Loading the Dataset

For the first part of our lab, we'll be using the GPA dataset we saw in class.

You can download the dataset by following to following link to the github storing the dataset: https://raw.githubusercontent.com/wadefagen/datasets/master/gpa/uiuc-gpa-dataset.csv

To download this file, you will want to **right click** and click **Save As**.

<!-- grade_id: cell-51fef9428e349da7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import pandas as pd
df = pd.read_csv('uiuc-gpa-dataset.csv')
### END SOLUTION
```

## Part 1 | Finding Total GPAs

Add two new columns, `Total Students` and `Average GPA`, that compute the total number of students in each course and the average weighted GPA for each course.

You should not include students with a `W` in either column and feel free to add other columns to help you out!

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python


df["Total Students"] = ...
df["Average GPA"] = ...



### BEGIN SOLUTION
df["Total Students"] = \
  df["A+"] + df["A"] + df["A-"] + \
  df["B+"] + df["B"] + df["B-"] + \
  df["C+"] + df["C"] + df["C-"] + \
  df["D+"] + df["D"] + df["D-"] + \
  df["F"]

df["Sum GPA"] = \
  (df["A+"] * 4) + (df["A"] * 4) + (df["A-"] * 3.67) + \
  (df["B+"] * 3.33) + (df["B"] * 3) + (df["B-"] * 2.67) + \
  (df["C+"] * 2.33) + (df["C"] * 2) + (df["C-"] * 1.67) + \
  (df["D+"] * 1.33) + (df["D"] * 1) + (df["D-"] * 0.67)

df["Average GPA"] = df["Sum GPA"] / df["Total Students"]
df
### END SOLUTION
```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
0      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
1      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
2      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
3      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
4      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
5      2018    Fall  2018-fa     AAS     105        Intro to Arab Am Studies   
6      2018    Fall  2018-fa     AAS     211    Asian Americans and the Arts   
7      2018    Fall  2018-fa     AAS     246    Asian American Youth in Film   
8      2018    Fall  2018-fa     AAS     287        Food and Asian Americans   
9      2018    Fall  2018-fa     ABE     100  Intro Agric & Biological Engrg   
10     2018    Fall  2018-fa     ABE     223    ABE Principles: Machine Syst   
11     2018    Fall  2018-fa     ABE     224   
```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(df[
  (df["YearTerm"] == "2010-su") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 410) \
]["Total Students"].sum() == 31), "It appears your \"Total Students\" column is incorrect."

assert(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "CS") & \
  (df["Number"] == 225) \
]["Total Students"].sum() == 697), "It appears your \"Total Students\" column is incorrect."

assert( abs(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 100) \
]["Average GPA"].mean() - 3.247) < 0.01), "It appears your \"Average GPA\" column is incorrect."
```

## Part 2 | Finding the easiest and hardest courses

### Puzzle 2.1: Hardest Courses in the Dataset

Using functions you already know, find the **30 hardest courses** at Illinois based on their "Average GPA".

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hardestCourses = ...
df_hardestCourses

### BEGIN SOLUTION
df_hardestCourses = df.nsmallest(30, 'Average GPA')
df_hardestCourses
### END SOLUTION
```

Output:
```
Year    Term YearTerm Subject  Number                   Course Title  \
1725   2018    Fall  2018-fa     LER     120    Contemporary Labor Problems   
14173  2016  Spring  2016-sp     LER     100  Introduction to Labor Studies   
21986  2014    Fall  2014-fa     LER     120    Contemporary Labor Problems   
24254  2014  Spring  2014-sp     LER     120    Contemporary Labor Problems   
27051  2013    Fall  2013-fa    MATH       2           Introductory Algebra   
24329  2014  Spring  2014-sp    MATH     220                       Calculus   
11901  2016    Fall  2016-fa     LER     120    Contemporary Labor Problems   
12169  2016    Fall  2016-fa     MCB     245     Human Anat & Physiol Lab I   
22033  2014    Fall  2014-fa     LLS     296  Bilingual US Latina/o Culture   
22197  2014    Fall  2014-fa     MCB     245     Human Anat & Physiol Lab I   
9102   2017  Spring  2017-sp     LER     100  Introduction to Labor Studies   
3076   2018  Spring  2018-sp    CHEM     101         Introd
```

### Puzzle 2.2: Hardest Courses in 2018

Find the **30 hardest courses in 2018** at Illinois based on their "Average GPA".

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_hardestCourses_2018 = ...
df_hardestCourses_2018

### BEGIN SOLUTION
df_hardestCourses_2018 = df[ df["Year"] == 2018 ].nsmallest(30, 'Average GPA')
df_hardestCourses_2018
### END SOLUTION
```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
1725  2018    Fall  2018-fa     LER     120     Contemporary Labor Problems   
3076  2018  Spring  2018-sp    CHEM     101          Introductory Chemistry   
1935  2018    Fall  2018-fa     MCB     245      Human Anat & Physiol Lab I   
4168  2018  Spring  2018-sp    MATH     446       Applied Complex Variables   
4871  2018  Summer  2018-su    CHEM     102             General Chemistry I   
4095  2018  Spring  2018-sp    MATH     112                         Algebra   
1723  2018    Fall  2018-fa     LER     100   Introduction to Labor Studies   
633   2018    Fall  2018-fa    CHEM     104            General Chemistry II   
379   2018    Fall  2018-fa    ASTR     404            Stellar Astrophysics   
1238  2018    Fall  2018-fa     FIN     300               Financial Markets   
4232  2018  Spring  2018-sp     MCB     250              Molecular Genetics   
1928  2018    Fall  2018-fa     MCB     245      Human Ana
```

### Puzzle 2.3: Easiest Courses in 2018

Find the **30 easiest courses in 2018** at Illinois based on their "Average GPA".

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easiestCourses_2018 = ...
df_easiestCourses_2018

### BEGIN SOLUTION
df_easiestCourses_2018 = df[ df["Year"] == 2018 ].nlargest(30, 'Average GPA')
df_easiestCourses_2018
### END SOLUTION
```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
3688  2018  Spring  2018-sp    EPSY     490  Prep for Recruitment Counselor   
1124  2018    Fall  2018-fa    EDUC     101   Education Orientation Seminar   
129   2018    Fall  2018-fa    ACCY     518       Financial Statement Fraud   
2044  2018    Fall  2018-fa     MUS     269                  Wind Orchestra   
2934  2018  Spring  2018-sp    BADM     420   Advanced Marketing Management   
1486  2018    Fall  2018-fa      IB     199      Undergraduate Open Seminar   
2436  2018    Fall  2018-fa    SPED     405     Gen Educator's Role in SPED   
4914  2018  Summer  2018-su    EPSY     408    Learning & Hum Dev w/ EdTech   
2792  2018  Spring  2018-sp    ANSC     103       Working With Farm Animals   
2998  2018  Spring  2018-sp    BIOE     436                Senior Design II   
1662  2018    Fall  2018-fa     LAS     101                Freshman Seminar   
860   2018    Fall  2018-fa      CS     498               
```

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_hardestCourses) == 30 ), "You should only have the 30 hardest courses."
assert( abs(df_hardestCourses["Average GPA"].mean() - 1.58) < 0.01 ), "Are you sure you have the 30 hardest courses?"

assert( len(df_hardestCourses_2018) == 30 ), "You should only have the 30 hardest courses in 2018."
assert( abs(df_hardestCourses_2018["Average GPA"].mean() - 2.07) < 0.01 ), "Are you sure you have the 30 hardest courses of 2018?"

assert( len(df_easiestCourses_2018) == 30 ), "You should only have the 30 easiest courses in 2018."
assert( abs(df_easiestCourses_2018["Average GPA"].mean() - 3.97) < 0.01 ), "Are you sure you have the 30 easiest courses of 2018?"
```

## Puzzle 3: Finding Data on  Every Subject

### Puzzle 3.1: Total Grades Given in Each Subject

Using `groupby` and `agg`, find the total number of each grade (eg: "A+"s, "A"s, "A-"s, etc.) for every subject in the DataFrame `df_grouped_subject`.

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python



df_grouped_subject = ...
df_grouped_subject

### BEGIN SOLUTION
g = df.groupby('Subject')
df_grouped_subject = g.agg('sum').reset_index()
df_grouped_subject
### END SOLUTION
```

Output:
```
Subject     Year  Number     A+      A     A-     B+      B    B-    C+  \
0       AAS   298016   19199    571   1726    714    460    422   198   116   
1       ABE   229611   38135    254   2200    721    632    827   210   116   
2      ACCY  3798626  689840   5437  18283  12947  12596  11368  4783  2267   
3       ACE  1492356  200528   2207  10864   4329   3698   5811  2264  1403   
4      ACES   245691   14066    569   1708    215     85    235    56    58   
5       ADV   698959  132912   1991   8075   3814   2826   4608  1260   806   
6        AE   570044   96822   1470   4726   2581   1803   2309   991   632   
7      AFRO   193325   19081    305    995    494    288    341   186   106   
8      AFST    20136    2172     58    160     93     42     28    13     4   
9      AGCM    52378    6805     40    309    103     67     74    19     5   
10     AGED   249725   40727    931   2332    957    569    467   212   132   
11      AHS    54360    7096    142    708     37     21
```

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_grouped_subject) == 161 ), "Are you sure you grouped by subjects?"
assert( df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["A"].sum() == 35107 ), "Are you sure you counted the right number of each subject?"
```

### Puzzle 3.2: Overall Average GPA

In `df_grouped_subject`, add a new column called `"Overall Average GPA"` that computes the total overall GPA of each subject at Illinois.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python
df_grouped_subject["Overall Average GPA"] = ...
df_grouped_subject

### BEGIN SOLUTION
df_grouped_subject["Overall Average GPA"] = df_grouped_subject["Sum GPA"] / df_grouped_subject["Total Students"]
df_grouped_subject
### END SOLUTION
```

Output:
```
Subject     Year  Number     A+      A     A-     B+      B    B-    C+  \
0       AAS   298016   19199    571   1726    714    460    422   198   116   
1       ABE   229611   38135    254   2200    721    632    827   210   116   
2      ACCY  3798626  689840   5437  18283  12947  12596  11368  4783  2267   
3       ACE  1492356  200528   2207  10864   4329   3698   5811  2264  1403   
4      ACES   245691   14066    569   1708    215     85    235    56    58   
5       ADV   698959  132912   1991   8075   3814   2826   4608  1260   806   
6        AE   570044   96822   1470   4726   2581   1803   2309   991   632   
7      AFRO   193325   19081    305    995    494    288    341   186   106   
8      AFST    20136    2172     58    160     93     42     28    13     4   
9      AGCM    52378    6805     40    309    103     67     74    19     5   
10     AGED   249725   40727    931   2332    957    569    467   212   132   
11      AHS    54360    7096    142    708     37     21
```

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( abs(df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["Overall Average GPA"].sum() - 3.236) < 0.01 ), "Your overall average GPA does not appear correct."
```

### Puzzle 3.3: Find the 15 easiest subjects at Illinois

Finally, find the 15 easiest subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-cb713763b3781870, points: 0, grade: false, solution: true -->
```python
df_easiest_subjects = ...
df_easiest_subjects

### BEGIN SOLUTION
df_easiest_subjects = df_grouped_subject.nlargest( 15, 'Overall Average GPA' )
df_easiest_subjects
### END SOLUTION
```

Output:
```
Subject    Year  Number    A+     A    A-   B+    B   B-  C+  \
34       CB    4022    1250     0   140     0    0   17    0   0   
40      CHP   12102    1194    42    94    15    8    5    0   0   
155     UKR    8048     452     5   190    43   20    8    4   0   
128    REES    6047     603    32    34     7    2    5    2   0   
53     EDUC  106864    8974   307  1688   220  102   92   23  16   
33      BUS  493460   26919   493  4525   924  413  230   82  31   
19     ARTE    8063    1079     9    58    10    5    4    3   0   
137     SBC   18161    4542   124    99    41   19    9    6   2   
59      EOL   82576   21905    62   866   188  101   69    8   2   
126     PSM   16126    4013   121    37    13    4   10    4   2   
11      AHS   54360    7096   142   708    37   21   52    4   0   
89       IS  829740  214129  1211  6353  1969  774  535  159  40   
41       CI  936346  202539  1381  7521  2000  912  630  173  76   
60     EPOL    6053    1650     1    94     9    2  
```

### Puzzle 3.4: Find the hardest 15 subjects at Illinois

Find the hardest 15 subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-eb5a80d62e1c7a45, points: 0, grade: false, solution: true -->
```python
df_hardest_subjects = ...
df_hardest_subjects

### BEGIN SOLUTION
df_hardest_subjects = df_grouped_subject.nsmallest( 15, 'Overall Average GPA' )
df_hardest_subjects
### END SOLUTION
```

Output:
```
Subject     Year  Number     A+      A     A-     B+      B     B-     C+  \
159      VM   161151   48575      0   1503      8      0   5029      1      0   
150     TAM   521580   83139   1611   3910   3026   3022   3722   2472   1982   
105    MATH  3951142  606935   9720  25074  16157  16442  22554  13336  11478   
121    PHYS  1409784  202288  11513  13751   9524   9157   9720   8607   7085   
99      LAT    12076     613     14     20     17     16     21      7     14   
136    SAME     6039     437      2     29     30     20     31     12     14   
36     CHBE   425030   88102    511   2952   1487   1565   2895   1410   1136   
38     CHIN     4023     445      0      8      9      4      9      3      1   
37     CHEM  3126162  282872   9759  38674  15498  15847  21077  12242  11048   
107     MCB  3256457  414467   5797  11913   9339   9153  11042   6904   5166   
125      PS  1081252  145728   1748   6582   5011   4236   5534   2872   1835   
51      ECE  2934499  505822   6
```

### Puzzle 3.5: Find your major!

Finally, find your major/subject in the DataSet and check out how hard it is!

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python

```

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. After you have saved and checkpointed, exit this notebook, and return to the Data Science DISCOVERY page on how to use git to turn this notebook into the course!

```python

```