# EC Notebook for Lecture 3: Ideal Experimental Design + Row Selection

This extra credit Python notebook will let you practice the material you saw in lecture!  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107!

This notebook is due at 11:30am on **Wednesday, Sept. 4** *(30 minutes before the next STAT 107 lecture)*.

## 1. Import the `pandas` library as `pd`

Using the code you learned in class, import the pandas library:

<!-- grade_id: cell-1be8ed504370241a, points: 0, grade: false, solution: true -->
```python
...

```

<!-- grade_id: cell-c481f083b4feb329, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

# Test #1:
# - This uses the `pd` (pandas) library to create a DataFrame.  This line will fail if `pd` has not been imported:
result = pd.DataFrame()
```

## 2. Load the "Hello Dataset"

As part of the first lecture, we created a dataset for this course!

1. The dataset can be found here: https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0


2. Download the dataset as a CSV:

  * In lecture you learned that Comma-Separated Values (CSV) files are an easy format for Data Science tools like Python to read.
  * Use the **File -> Download** menu to choose the format you want to download your data in from Google Sheets.
  * Once downloaded, move the file from your "Downloads" to the same directory as this Python Notebook.  If you use the recommended file location, the location will be: **Your Desktop -> stat107 -> netid -> ec3**.


3. Run the next cell to load the dataset into Python into the variable name `df`. 

  * The file name you are loading must **EXACTLY MATCH** the file name on your computer, including capitalization.  You will need to edit the code (or the downloaded file name) so the file name matches.

<!-- grade_id: cell-92f9663de5efb2a7, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-e93df3b99bed5c14, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

# Ensure `Wade` exists somewhere in the data:
assert(df[ df.Name == 'Wade' ].Name.values[0] == 'Wade'), "No row is found that contains `Wade`.  Did you import the right dataset?"

# Ensure `Karle` exists somewhere in the data:
assert(df[ df.Name == 'Karle' ].Name.values[0] == 'Karle'), "No row is found that contains `Karle`.  Did you import the right dataset?"
```

## 3. Extra Credit Problems

### Puzzle 3.1: A Data Science pill

In lecture, you learned that Karle and Wade created a new pill that they think will help make students EXPERT Python programmers.  Before distributing this pill, they need to do an experiment to see if it actually works.  Assume everyone in the class has given permission to be involved in the experiment, but we can only accommodate **80** people.  Use Python to take a random sample of 80 students from the class and store it in the variable named `df_80`.

<!-- grade_id: cell-eb21607130c75ec7, points: 0, grade: false, solution: true -->
```python
df_80 = ...
df_80


```

<!-- grade_id: cell-ca892fdb27a8f534, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.


# Ensure df_80 has exactly 80 observations (rows)
assert(len(df_80) == 80), 'The variable `df_80` does not have exactly 80 observations/rows (it has ' + str(len(df_80)) + ' rows).'

# Ensure df_80 does not contain only the first 80 rows
full_match = True
for i in range(0, 80):
    if df_80.index[i] != i:
        full_match = False
        break
        
assert(full_match == False), 'The variable `df_80` contains the FIRST 80 observations/rows.  Make sure to make a RANDOM selection of the 80 rows.'
```

### Puzzle 3.2: Finding Ten Winners

Suppose a contest is held in Data Science Discovery and we need to pick **ten (10) winners at random**.  Use Python to find 10 winners at random and store those winners in the variable `winners`.

<!-- grade_id: cell-40bda3e3d35ff1b0, points: 0, grade: false, solution: true -->
```python
winners = ...
winners


```

<!-- grade_id: cell-0273f0f917db94e1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

# Ensure `winners` has exactly 10 observations (rows)
assert(len(winners) == 10), 'The variable `winners` does not have exactly 10 observations/rows (it has ' + str(len(winners)) + ' rows).'

# Ensure `winners` does not contain only the first 10 rows
full_match = True
for i in range(0, 10):
    if winners.index[i] != i:
        full_match = False
        break
        
assert(full_match == False), 'The variable `winners` contains the FIRST 10 observations/rows.  Make sure to make a RANDOM selection of the 10 rows.'
```

### Puzzle 3.3:

Take a look at the list of winners from the previous puzzle and **re-run** the cell (Shift+Enter).  Double click this cell 
to answer the next question:

- Do the winners change each time you run it?

## 4. Submit Your Work

In this notebook:

1. Click **File** -> **Save and Checkpoint** (to save your work)
2. Click **File** -> **Close and Halt** (to exit this notebook)

Follow the instructions on the STAT 107 website to submit your work.