# EC Notebook for Lecture 22: Binomial Distribution and its Expectation

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Oct. 21** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## Simulation of Binomial Distribution

In the lecture, we learned about the binomial distribution. First, we will try to simulate an example with binomial distribution and plot the results. Next, we will compare the true expectation with our result.

## Question 1: Writing a simulation

Suppose, we have a bag filled with **3 Red marbles, 4 Blue marbles and 3 Green marbles**. One by one, we will draw **20** marbles with replacement (i.e, after each draw, we put back the marble in the bag). We are interested in counting the **number of Red marbles** that we see in our 20 draws.  

### Part 1

First, write a function (`red_count()`)that counts the number of Red marbles in the 20 draws.

<!-- grade_id: cell-a0315ff1f3304a75, points: 0, grade: false, solution: true -->
```python
def red_count():
    return 0


```

<!-- grade_id: cell-f5054d49f276f230, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(0 < red_count() < 20), "The count of Red marbles may be incorrect"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

### Part 2: Storing results in a dataframe

Now, write a simulation to run this experiment (drawing 20 marbles with replacement) 1000 times. Store your results in the dataframe named `df`. And print out a few rows from `df`.

<!-- grade_id: cell-5cf0a806a71f1906, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = ...


# Print out a few random rows:
print(df.sample(5))
```

<!-- grade_id: cell-e492e1c42fb29fc7, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(df) >= 1000), "Make sure that you have run the simulation at least 1000 times"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Part 3: Plot a histogram

Plot a histogram of your results.

<!-- grade_id: cell-5dce6e0d379b8201, points: 0, grade: false, solution: true -->
```python


```

## Question 2:

Using the formula that you learned in class today, calculate the true expected value of number of red marbles that we will see in 20 draws. 

<!-- grade_id: cell-5cfbe9d3dc0acb8b, points: 0, grade: false, solution: true -->
```python
expected_red_count = ...
expected_red_count 


```

<!-- grade_id: cell-3a270c0cb136d5a0, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(expected_red_count*1000 == 6000), "Make sure that you have used the correct formula for expected value for Binomial Distribution"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

##### Does the true expected value match with the histogram based on your results. If not, what can we change to achieve that!? (You don't have to answer it. Try changing something in your code)