# EC Notebook for Lecture 21: Random Variables

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Friday, Oct. 18** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## What's the random variable

In the lecture, we learned the definition of **random variable**.  In this section, we will get familiar with random variables by doing some practices to generate the random variables.

### Question 1:

In Lecture, Prof. Karle showed off how to set up a random variable $X$ for a game of roulette:

> A roulette table has 38 slots—18 red, 18 black and 2 green. If you bet \\$1 on red and it comes up red, then you win \\$1; otherwise, you lose \\$1.  Suppose you simulate playing roulette 50 times, betting \\$1 each time.

Write a simulation of spinning the roulette wheel at 50 times and record your total profit (or loss).  Run that simulation at least 1,000 times.

<!-- grade_id: cell-5ced175279b31f80, points: 0, grade: false, solution: true -->
```python
# Write the simulation of spinning the roulette wheel at 50 times:



```

<!-- grade_id: cell-d9db332955eef16d, points: 0, grade: false, solution: true -->
```python
# Create a DataFrame called `df` of your total profit for each simulation.
# - Make sure to store your result in the DataFrame column `profit`




```

### Quesiton 1b: Create a hisogram of your winnings

Create a histogram of your DataFrame to show the histogram of your winnings (or loss).

<!-- grade_id: cell-a77c7500aa273a3e, points: 0, grade: false, solution: true -->
```python



```

<!-- grade_id: cell-e3b938d515a724f4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert( len(df) >= 1000 ), "Your DataFrame must contain at least 1,000 simulations of roulette"
assert( "profit" in df ), "Your DataFrame must contain the column `profit`"
assert( len(df[ df.profit == 0]) > 0 ), "Your DataFrame should contain some games with 0 profit"
assert( len(df[ df.profit < 0]) > 0 ), "Your DataFrame should contain some games with negative profit"
assert( len(df[ df.profit > 0]) > 0 ), "Your DataFrame should contain some games with positive profit"
assert( len(df[ df.profit < 0 ]) > len(df[ df.profit > 0 ]) ), "Your DataFrame seems suggest you would make money playing roulette -- is your simulation correct?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print(f"- Your average simulation profit is -${ df['profit'].mean() * -1 } per 50 games played" )
```

### Question 2: Discrete random variable

In the lecture, we have the definition of discrete random variables, and we will use this definition to do an simulation.

A box includes 81 balls with four different colors:
- 15 green,
- 24 red,
- 31 blue, and
- 11 pink balls

Wade is trying to pick up one ball from the box for each time. And if waf picks up the pink ball **or** red ball, he will win a GIANT Teddy Bear -- otherwise nothing.

Suppose Prof. Wade repeats this game 250 times.  Simulate (just once) how many teddy bears will he now own after 250 times playing this game.

<!-- grade_id: cell-a0315ff1f3304a75, points: 0, grade: false, solution: true -->
```python
teddy_bears = ...
teddy_bears


```

<!-- grade_id: cell-f5054d49f276f230, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(50 < count_100 < 150), "The count of Teddy Bear may be incorrect"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

### Question 3: Expectation and variance of discrete random variables

In the lecture, we learned the formulas to calculate the expectation and variance of the discrete random variables. Please use the following distribution information to calculate the expectation and variance of this discrete random variable:

| X | P(X) |
|---|------|
| 2 |  10% |
| 4 |  25% |
| 6 |  15% |
| 8 |  30% |
|10 |  10% |
|12 |  10% |

<!-- grade_id: cell-5cf0a806a71f1906, points: 0, grade: false, solution: true -->
```python
EX = ...
VarX = ...


```

<!-- grade_id: cell-e492e1c42fb29fc7, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(6.6 < EX < 6.8), "The expectation value is incorrect"
assert(8.50 < VarX < 8.52), "The variance is incorrect"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```