## EC Notebook for Lecture 26: Confidence Intervals

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Oct. 30** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Importing pandas and reading data
For this extra credit notebook, we will still use the GPA dataset to do some exercise. At first, you need to import the library pandas and reading the dataset.

### Calculate the size of each class
After loading the data, calculate the size of each class and add an extra column named `size` to the DataFrame. (*Hint: students with a W grade is also a part of the class.*) Observe the distribution of class size.

<!-- grade_id: cell-3182d2cfb9079fe2, points: 0, grade: false, solution: true -->
```python
#import pandas
df = ...


df["size"].hist(bins=20)
```

### Class Sampling

Now lets create a function, which will sample `n` sections from the DataFrame and calculate the following values:

1. the average size of samples
2. the standard error of samples

and return them.

<!-- grade_id: cell-060c81a4477630d2, points: 0, grade: false, solution: true -->
```python
def sample_class(n):
    ...
    return (sample_mean, sample_sem)


```

### Distribution of Average
Write a simulation to run the above function for 1000 times with `n=100` and store the results you get for each time in a DataFrame called `df_class`, which contains two cloumns `mean` and `sem`.

<!-- grade_id: cell-8b0f0ef67575550d, points: 0, grade: false, solution: true -->
```python
df_class = ...



# Let’s build visualization:
df_class["mean"].hist(bins=20)
```

###  Confidance Interval

What we have done above shows that when you sample averages from a non-normal distribution, you get a normal distribution out. Now let's calculate the **95%** confidance interval bound given the above sampling.

<!-- grade_id: cell-9b6af2e58eb5bdd7, points: 0, grade: false, solution: true -->
```python
df_class["left_ci"] = ...
df_class["right_ci"] = ...


df_class.sample(5)
```

<!-- grade_id: cell-621e2757b30ebf38, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1,2,3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df["size"]) == len(df)), "Make sure you added the correct size of each class."
assert(len(df_class["mean"]) == 1000), "Make sure you create the right dataframe df_class."
assert(len(df_class["sem"]) == 1000), "Make sure you create the right dataframe df_class."
assert(len(df_class["left_ci"]) == 1000), "Make sure you create the right dataframe df_class."
assert(len(df_class["right_ci"]) == 1000), "Make sure you create the right dataframe df_class."
assert(abs((df_class["left_ci"] +  df_class["right_ci"]- 2*df_class["mean"]).sum()) < 0.0001), "The bound seems not to be right."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

```python

```

```python

```