# EC Notebook for Lecture 4: Blocking and Conditionals

This extra credit Python notebook will let you practice the material you saw in lecture!  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107!

This notebook is worth +1 if turned in before 11:30am on **Friday, Sept. 6** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will continue to use the "Hello" dataset introduced last lecture!

Using your @illinois.edu address, download the dataset as a CSV and **make sure to save or move it to the same directory as this notebook**: https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0

Use the following cell to load the pandas library and read the DataSet as a DataFrame:

<!-- grade_id: cell-7543f21c9c406c26, points: 0, grade: false, solution: true -->
```python
...

```

## 2. Finising the Lecture Problem

In lecture, we left you with a final puzzle.  From the lecture handout:

> Use Python to create a control group of 20 students, using blocking to control for the influence we suspect having more than 1 siblings might have.  Make sure you have an equal number of students with 0 or 1 sibling and an equal number of students with 2 or more siblings.

We will complete this in four separate steps to get practice on each part.

#### 2.1: Sibling Block #1

Creating a DataFrame that contains all students with 0 or 1 siblings and call it `block1`:

<!-- grade_id: cell-42b76d09e5d1ea90, points: 0, grade: false, solution: true -->
```python
block1 = ...
block1

```

#### 2.2: Sibling Block #2

Create a DataFrame that contains all students with 2 or more siblings and call it `block2`:

<!-- grade_id: cell-ad3eef9bc1500907, points: 0, grade: false, solution: true -->
```python
block2 = ...
block2

```

#### 2.3: Selecting equal groups from block1 and block2

**Randomly** select an equal number of students from each of `block1` and `block2` and store them in `control1` and `control2`.  We will combine both groups to make the full control group:

<!-- grade_id: cell-71025862479973ff, points: 0, grade: false, solution: true -->
```python
control1 = ...
control1

```

<!-- grade_id: cell-0a28afaa87e1114d, points: 0, grade: false, solution: true -->
```python
control2 = ...
control2

```

#### 2.4. Combine the two sub-groups into a single control group

Use the following code to merge the contents of `control1` and `control2` together into one large group called `control`:

<!-- grade_id: cell-f5b1dce206b67c5a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-78cde41a46e52f67, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(control) == 20 ), "Your control group must only have 20 courses in it."
assert( len(control[control.Siblings < 2] ) == 10 ), "Your control group must only have 10 people with 0-1 siblings in it."
assert( len(control[control.Siblings >= 2] ) == 10 ), "Your control group must only have 10 people with 2+ siblings in it."
```

## 3. Create a treatment group

Let's extend what you did onto a new topic.  Using this week's dataset, the course catalog at Illinois, use Python to create a treatment group of 10 courses, using blocking to control for the influence we suspect being an upper-level course might have.

An upper-level course is a course that has a number course number of 300 or higher.  Make sure you have an equal number of lower-level and upper-level courses in your final treatment group.

You can find the course catalog dataset here: https://docs.google.com/spreadsheets/d/1ZnpZpipyui_zIIBblmalRYiyEG1glUKinqJtQxVL04o/edit#gid=0

You have several cells to work in if you need them.  The only requirement is that your final treatment group must be randomly selected and be stored in the variable `treatment`.

<!-- grade_id: cell-3c72e431e5e37750, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-ddcd363538ddaefb, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-600e3e734fc2a695, points: 0, grade: false, solution: true -->
```python
# Make sure your final answer is in `treatment`:
treatment = ...
treatment


```

<!-- grade_id: cell-1b422f5e6c77b87e, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(treatment) == 10 ), "Your treatment group must only have 10 courses in it."
assert( len(treatment[treatment.Number < 300] ) == 5 ), "Your treatment group must only have 5 lower-level courses in it."
assert( len(treatment[treatment.Number >= 300] ) == 5 ), "Your treatment group must only have 5 upper-level courses in it."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)