# EC Notebook for Lecture 7: Working with columns

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Friday, Sept. 13** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Hello" dataset -- as we did in lecture -- to answer a few questions on working with columns.

Using your @illinois.edu address, you can access the "Hello" dataset here: https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0

<!-- grade_id: cell-0d5f5f22a56cb05c, points: 0, grade: false, solution: true -->
```python

```

# 2. Do upperclass students study less?

Do seniors study less than freshman?  Using the "Hello" dataset, find the average number of hours freshman study (`freshman_avg_study_hours`) and the average number of hours seniors study (`senior_avg_study_hours`) each week:  

<!-- grade_id: cell-0c5e6b97fdf25f9a, points: 0, grade: false, solution: true -->
```python


freshman_avg_study_hours = ...
freshman_avg_study_hours


```

<!-- grade_id: cell-ba360e51c3559faa, points: 0, grade: false, solution: true -->
```python


senior_avg_study_hours = ...
senior_avg_study_hours


```

<!-- grade_id: cell-e03adddf03d4c81a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(freshman_avg_study_hours == 3.73), "The average number of hours studying for freshman appears incorrect."
assert(senior_avg_study_hours - 3.77 < 0.01), "The average number of hours studying for seniors appears incorrect."
```

# 3. Do introverts text less than extroverts?

In our dataset, you responded with the number of texts recieved yesterday.  Find the average number of texts sent by introverts, by extroverts, and by ambiverts:

<!-- grade_id: cell-0ded48531f57011e, points: 0, grade: false, solution: true -->
```python
introverts_avg_txts = ...
introverts_avg_txts


```

<!-- grade_id: cell-3cadfc178d2b60a5, points: 0, grade: false, solution: true -->
```python
ambiverts_avg_txts = ...
ambiverts_avg_txts


```

<!-- grade_id: cell-975d7137d0ca4494, points: 0, grade: false, solution: true -->
```python
extroverts_avg_txts = ...
extroverts_avg_txts


```

<!-- grade_id: cell-85b1350f5048af0d, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(introverts_avg_txts < extroverts_avg_txts), "This does not appear correct -- check your introvert and extrovert logic."
assert(introverts_avg_txts < ambiverts_avg_txts), "This does not appear correct -- check your introvert and ambivert logic."
assert(extroverts_avg_txts < ambiverts_avg_txts), "This does not appear correct -- check your extrovert and ambivert logic."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)