# EC Notebook for Lecture 12: Quartiles and Box Plots

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30 am on **Wednesday, Sept. 25** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "words" dataset -- as we did in lecture 12 -- to visually present quantitative data.

Using your @illinois.edu address, you can access the "words" dataset here: 
https://docs.google.com/spreadsheets/d/18zQc-cYb376QA478lgWnMbQj39M9wVqHtjbvLP0PTSc/edit#gid=0

Always, remember to load your data first.

<!-- grade_id: cell-1628b09530950d23, points: 0, grade: false, solution: true -->
```python

```

## 2 Box Plots and Quartiles
Today we will observe and analyze our perception of words that describe uncertainty. Firstly, let's draw the **boxplot** for the words "*We Believe*" and "*We Doubt*" using the function you learned from the lecture. 

### Question 1:
Draw two box plots for columns "*We Believe*" and "*We Doubt*", and then compute the median of them.


<!-- grade_id: cell-99d959c1fa5c1922, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
plt.figure(figsize=(8, 1.5))
median_believe = ...
median_doubt = ...
print(median_believe)
print(median_doubt)
#Write the codes for the boxplot



```

<!-- grade_id: cell-6d95af9e6dbf7763, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(median_believe == df["We Believe"].quantile(0.5)), "The median of believe does not appear correct."
assert(median_doubt == df["We Doubt"].quantile(0.5)), "The median of doubt does not appear correct ."
```

You have already known that median is **Q2** among Quartiles and how to get it. But how about **Q1** and **Q3**? In pandas, there is a function called [pandas.quantile](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.quantile.html) we can utilize.

### Question 2:
Following the given example, compute **Q1**, **Q3** and **Interquartile Range (IQR)** for "We Believe" and "We Doubt".

<!-- grade_id: cell-0f2b9437c5b3e29a, points: 0, grade: false, solution: true -->
```python
# Example of pandas.quantile
# you can check whether they are the same as the medians you just computed in the previous cell
Q2_believe = df["We Believe"].quantile(0.5)
Q2_doubt = df["We Doubt"].quantile(0.5)
print(Q2_believe)
print(Q2_doubt)

# Calculate the Q1 and Q3 by youself! 
# Hints: you can directly use 0.25 or 0.75 as parameter to request for Q1 or Q3
Q1_believe = ...
Q1_doubt = ...
Q3_believe = ...
Q3_doubt = ...
IQR_believe = ...
IQR_doubt = ...
print(IQR_believe)
print(IQR_doubt)


```

<!-- grade_id: cell-6f4ebbb95f35fe62, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(Q1_believe == 55), "The Q1 of believe does not appear correct."
assert(Q3_believe == 80), "The Q3 of believe does not appear correct."
assert(Q1_doubt == 10), "The Q1 of doubt does not appear correct."
assert(Q3_doubt == 35), "The Q3 of doubt does not appear correct."
assert(IQR_believe == 25), "The IQR of believe does not appear correct."
assert(IQR_doubt == 25), "The IQR of doubt does not appear correct."
```

Now you have calculated a lot of useful information! Try to find out outliers of "We Doubt" and "We Believe" using all the above results you got! Remember to observe the box plot carefully to determine the kind of outlier.

### Question 3 
Calculate the number of **outliers** of "We Doubt" and "We Believe" column.

<!-- grade_id: cell-ea88dcf40cc90ada, points: 0, grade: false, solution: true -->
```python
outliers_believe = ...
outliers_doubt = ...
print(outliers_believe)
print(outliers_doubt)


```

<!-- grade_id: cell-357bac434e050bca, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(int(outliers_believe) == 6), "The outliers of believe does not appear correct."
assert(int(outliers_doubt) == 4), "The outliers of doubt does not appear correct."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)

```python

```