# EC Notebook for Lecture 5: Combining DataFrames

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Monday, Sept. 9** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Course Catalog" dataset that is our dataset of the week.

Here's the link: https://docs.google.com/spreadsheets/d/1ZnpZpipyui_zIIBblmalRYiyEG1glUKinqJtQxVL04o/edit#gid=0

Use the following cell to load the pandas library and read the DataSet as a DataFrame:

<!-- grade_id: cell-ba4c939dda92ddf8, points: 0, grade: false, solution: true -->
```python
...

```

## 2. Selecting all SUBJ 100 and SUBJ 101

In many subjects, 100 and 101 courses serve as freshman entry-level courses.  Create single DataFrame called `entryCourses` that contains all courses with numbers `100` or `101`.  *There are MANY different ways you can do this, including using `df.merge` if needed.*

<!-- grade_id: cell-ca1503cca7507a80, points: 0, grade: false, solution: true -->
```python


entryCourses = ...
entryCourses

```

<!-- grade_id: cell-51c14f7e00a20667, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert( len(entryCourses) == 116 ), "`entryCourses` must contain all 116 entry-level courses."
assert( len(entryCourses[entryCourses.Number == 100] ) == 48 ), "`entryCourses` must contain all 48 courses as SUBJ `100`."
```

## 3. Creating a `control` and `treatment` group

Using the `entryCourses` you created in Puzzle 2, create a `treatment` and `control` group that each contain 50 random courses to use an experiment you are developing.

*Remember: A course cannot be in both a treatment and control group!*

<!-- grade_id: cell-8ba4be23607128e9, points: 0, grade: false, solution: true -->
```python
# Create the control group:


control = ...
control


```

<!-- grade_id: cell-e2fbfed754aeba62, points: 0, grade: false, solution: true -->
```python
# Create the treatment group:


treatment = ...
treatment


```

<!-- grade_id: cell-8d7d473f143588c7, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(control) == 50 ), "The control group must contain exactly 50 courses."
assert( len(treatment) == 50 ), "The control group must contain exactly 50 courses."

assert( len(control[control.Number > 101] ) == 0 ), "`control` must contain only entry-level courses."
assert( len(treatment[treatment.Number > 101] ) == 0 ), "`treatment` must contain only entry-level courses."

assert( len(pd.merge(control, treatment, how='inner', on=['Subject', 'Number'])) == 0 ), "Courses in `treatment` and `control` must be different."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)