# EC Notebook for Lecture 9: Grouping Data in Pandas

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Wednesday, Sept. 18** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Importing pandas and reading data
For this notebook, we will use the Illinois GPA Distribution dataset -- as we did in lecture -- to answer a few questions on working with columns.

You can access the "GPA" dataset here: https://github.com/wadefagen/datasets/tree/master/gpa

<!-- grade_id: cell-9aafb896bbc8e800, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = pd.read_csv("gpa.csv")
```

## 2. Review: Boolean Logic and Conditionals
The Illinois CS and STAT Students Centers want to calculate the number of students who did an excellent job in the courses of CS and STAT. In this case, they define those who got an "A+" or "A" are excellent. Also, considering the course difficulty, they think those who got an "A-" in advanced-level courses (course level >= 400) should be counted in excellent students. Can you help them to get the number?

### Question 1: 
Using the GPA dataset, calculate the number of students who got "A+" or "A" (in all level CS and STAT courses) or "A-" (only in advanced-level CS and STAT courses) in the DataFrame `df`. 

<!-- grade_id: cell-c0705f665f8b46dd, points: 0, grade: false, solution: true -->
```python
reg = df[(df["Subject"]=="CS") | (df["Subject"]=="STAT")]
adv = reg[reg["Number"]>=400]

reg_count = df.loc[df['Subject']=="CS", 'A+'].sum()
reg_count = reg_count+ df.loc[df['Subject']=="STAT", 'A+'].sum()
reg_count = reg_count+ df.loc[df['Subject']=="CS", 'A'].sum()
reg_count = reg_count+ df.loc[df['Subject']=="STAT", 'A'].sum()
reg_count = reg_count+ adv.loc[df['Subject']=="CS", 'A-'].sum()
reg_count = reg_count+ adv.loc[df['Subject']=="STAT", 'A-'].sum()

num_excellent_students = reg_count
```

<!-- grade_id: cell-98d118325ba78938, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(num_excellent_students == 87191), "The number does not appear correct."
```

## 3. Review: Adding a New Variable (column)
Thanks to your great job, CS and STAT Students Centers know better about their courses. And they recommend you to other departments' students center, so they're all coming to ask you for a favor. This time they are curious about the percentage of students who got "F" in each course.

### Question 2: 
Using the GPA dataset, calculate the number of students who got "F" of each course and add a new column titled `F rate`. To test your correctness, print the average `F rate` of all the courses.

<!-- grade_id: cell-e0c17bd27629d741, points: 0, grade: false, solution: true -->
```python
for i in range(46544):
    if (df["A+"][i+1] + df["A"][i+1] + df["A-"][i+1] + df["B+"][i+1] + df["B"][i+1] + df["B-"][i+1] + df["C+"][i+1] + df["C"][i+1] + df["C-"][i+1] + df["D+"][i+1] + df["D"][i+1] + df["D-"][i+1] + df["F"][i+1] + df["W"][i+1]) != 0:
        df["F rate"][i+1] = df["F"][i+1] / df["A+"][i+1] + df["A"][i+1] + df["A-"][i+1] + df["B+"][i+1] + df["B"][i+1] + df["B-"][i+1] + df["C+"][i+1] + df["C"][i+1] + df["C-"][i+1] + df["D+"][i+1] + df["D"][i+1] + df["D-"][i+1] + df["F"][i+1] + df["W"][i+1]
    else:
        df["F rate"][i+1] = 0


avg_f_rate = df["F rate"].mean()
df["F rate"]
```

<!-- grade_id: cell-576f5fd03e1f7da2, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(avg_f_rate - 0.017564803426637177) < 0.01), "The rate does not appear correct."
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-28a440bac74e> in <module>()
      4 # - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
      5 
----> 6 assert(abs(avg_f_rate - 0.017564803426637177) < 0.01), "The rate does not appear correct."

AssertionError: The rate does not appear correct.
```